# BT01: Ẩn tin mật trên văn bản bằng phương pháp chèn khoảng trắng <font color="grey">(chèn khoảng trắng giữa các từ để làm cho văn bản được căn lề 2 bên)


Họ tên: Nguyễn Minh Tâm

MSSV: 1712746

---
## Cách làm bài và nộp bài
&#9889; Bạn lưu ý là mình sẽ dùng chương trình hỗ trợ chấm bài nên bạn cần phải tuân thủ chính xác qui định mà mình đặt ra, nếu không rõ thì hỏi, chứ không nên tự tiện làm theo ý của cá nhân.

**Cách làm bài**

Bạn sẽ làm trực tiếp trên file notebook này. Đầu tiên, bạn điền họ tên và MSSV vào phần đầu file ở bên trên. Trong file, bạn làm bài ở những chỗ có ghi là:
```python
# YOUR CODE HERE
raise NotImplementedError()
```
hoặc đối với những phần code không bắt buộc thì là:
```python
# YOUR CODE HERE (OPTION)
```
hoặc đối với markdown cell thì là:
```markdown
YOUR ANSWER HERE
```
Tất nhiên, khi làm thì bạn xóa dòng `raise NotImplementedError()` đi.
Đối những phần yêu cầu code thì thường ở ngay phía dưới sẽ có một (hoặc một số) cell chứa các bộ test để giúp bạn biết đã code đúng hay chưa; nếu chạy cell này không có lỗi gì thì có nghĩa là qua được các bộ test. Trong một số trường hợp, các bộ test có thể sẽ không đầy đủ; nghĩa là, nếu không qua được test thì là code sai, nhưng nếu qua được test thì chưa chắc đã đúng hoàn toàn.

Trong khi làm bài, bạn có thể cho in ra màn hình, tạo thêm các cell để test. Nhưng khi nộp bài thì bạn xóa các cell mà bạn tự tạo, xóa hoặc comment các câu lệnh in ra màn hình. Bạn lưu ý <font color=red>không được tự tiện xóa các cell hay sửa code của Thầy</font> (trừ những chỗ được phép sửa như đã nói ở trên).

Trong khi làm bài, thường xuyên `Ctrl + S` để lưu lại bài làm của bạn, tránh mất mát thông tin.


*Nên nhớ mục tiêu chính ở đây là <font color=green>học, học một cách chân thật</font>. Bạn có thể thảo luận ý tưởng với bạn khác, nhưng <font color=green>code và bài làm phải là của bạn, dựa trên sự hiểu thật sự của bạn</font>. <font color=red>Nếu vi phạm thì sẽ bị 0 điểm cho toàn bộ môn học.</font>*

**Cách nộp bài**

Khi chấm bài, đầu tiên mình sẽ chọn `Kernel` - `Restart Kernel & Run All Cells`, để restart và chạy tất cả các cell trong notebook của bạn; do đó, trước khi nộp bài, bạn nên chạy thử `Kernel` - `Restart Kernel & Run All Cells` để đảm bảo mọi chuyện diễn ra đúng như mong đợi.

Sau đó, bạn tạo thư mục nộp bài theo cấu trúc sau:
- Thư mục `MSSV` (vd, nếu bạn có MSSV là 1234567 thì bạn đặt tên thư mục là `1234567`)
    - Thư mục `BT01`
        - File `BT01-AnTinMatTrenVanBan_KhoangTrang.ipynb` (không cần nộp các file khác)

Cuối cùng, bạn nén thư mục `MSSV` này lại và nộp ở link trên moodle. <font color=red>Bạn lưu ý tuân thủ chính xác cấu trúc này.</font>

## 2. Import

```python
from bitarray import bitarray
```

## 3. Hàm nhúng (6đ)

<!-- grade_id: cell-2f21b430087c5c35, points: 0, grade: false, solution: true -->
```python
def embed(msg_file, cover_text_file, text_width, stego_text_file):
    '''
    Nhúng tin mật vào văn bản bằng phương pháp chèn khoảng trắng 
    để văn bản được căn lề 2 bên.
    
    Các tham số:
        msg_file (str): Tên file chứa secret message.
        cover_text_file (str): Tên file chứa cover text.
        text_width (int): Chiều dài của dòng sau khi được căn lề.
        stego_text_file (str): Tên file chứa stego text.
    Giá trị trả về:
        bool: True nếu nhúng thành không, False nếu không đủ chỗ để nhúng. 
    '''
    # Read cover text file
    with open(cover_text_file, 'r') as f:
        cover_text_lines = f.read().splitlines()
    
    # Read msg file
    with open(msg_file, 'r') as f:
        msg = f.read()
    
    # Convert msg to msg bits
    msg_bits = bitarray();
    msg_bits.frombytes(msg.encode('utf-8'))
    
    # Embed msg bits into cover text lines
    stego_text = ""
    b = 0
    ended = False # tức bit 1 trong dãy 10000 đã được add vào chưa
    
    for idx, line in enumerate(cover_text_lines):
        if(idx == len(cover_text_lines) - 1): # last line in file
            stego_text += line;
            break;
        elif(cover_text_lines[idx+1] == ""): # last line in paragraph
            stego_text += (line+"\n");
            continue;
        
        # support variables
        next_stego_line = "" # kết quả
        words = line.split(' ')
        n = line.count(' ') # n = số khoảng trắng của dòng l
        m = text_width - len(line) # m = số khoảng trắng cần phải chèn để dòng l được căn lề
        num_spaces_inserted = 0; # số khoảng cách đã duyệt
        n_bits = min(m, n-m) if min(m, n-m) > 0 else 0 # số msg_bit nhúng được = số cặp
        num_words_left = len(words) # số từ còn lại
        num_spaces_left = n+m # số khoảng cách còn lại
        
        if m < n and m > 0:
            # loop pair spaces
            for i in range(n_bits):
                if b < len(msg_bits):
                    #nhúng msg_bits[b]
                    if(msg_bits[b] == True): # bit 1 => 1 kc + 2 kc
                        next_stego_line += (words[i*2] + " " + words[i*2+1] + "  ")
                    else: # bit 0 => 2 kc + 1 kc
                        next_stego_line += (words[i*2] + "  " + words[i*2+1] + " ")
                    
                    b += 1;
                else:
                    if(ended): # bit 1 dãy 100.. đã add vào, thêm bit 0 vào
                        next_stego_line += (words[i*2] + "  " + words[i*2+1] + " ")
                    else: # bit 1 chưa thêm, thêm bit 1 vào
                        next_stego_line += (words[i*2] + " " + words[i*2+1] + "  ")
                        ended = True
                num_spaces_inserted += 3;
            num_words_left -= n_bits*2
            num_spaces_left -= num_spaces_inserted;
            
        # xử lý những từ hay khoảng cách còn lại
        if(num_words_left > 1): # không phải từ cuối
            for i in range(num_words_left):
                num_space = 0 # số khoảng cách cần thêm cho từ kế
                if(i != num_words_left - 1): #not last_word
                    num_space_each = int(num_spaces_left / (num_words_left-1));
                    num_space_addition = (1 if (i < (num_spaces_left-(num_words_left-1)*num_space_each)) else 0);
                    num_space = num_space_each + num_space_addition
                next_stego_line += (words[n_bits*2 + i] + " "*num_space); 
        else: # từ cuối cùng
            next_stego_line += words[n_bits*2]
        stego_text += (next_stego_line + "\n");
    if (not ended): # chưa nhúng được bit 1 trong dãy 100 thì nhúng thất bại
        return False
        
    # Write stego text to file
    with open(stego_text_file, 'w') as f:
        f.write(stego_text)   
    
    return True
        
        
```

<!-- grade_id: cell-057a58bf038e1e60, points: 5, grade: true, solution: false -->
```python
# TEST 1            
result = embed('msg.txt', 'cover.txt', 70, 'stego.txt')
assert result == True
with open('stego.txt', 'r') as f:
    stego_text = f.read()
with open('correct_stego.txt', 'r') as f:
    correct_stego_text = f.read()
assert stego_text == correct_stego_text
```

<!-- grade_id: cell-f625774497df6542, points: 1, grade: true, solution: false -->
```python
# TEST 2
result = embed('msg2.txt', 'cover.txt', 70, 'stego.txt')
assert result == False
```

## 4. Hàm rút trích (4đ)

<!-- grade_id: cell-f3f9b34770f282aa, points: 0, grade: false, solution: true -->
```python
def extract(stego_text_file, extr_msg_file):
    '''
    Hàm rút trích tin mật đã được nhúng bằng phương pháp chèn khoảng trắng.
    
    Các tham số:
        stego_text_file (str): Tên file chứa stego text.
        extr_msg_file (str): Tên file chứa secret message được rút trích.
    '''
    # Read stego text file
    with open(stego_text_file, 'r') as f:
        stego_text_lines = f.read().splitlines()
    
    # Extract msg bits
    extr_msg_bits = bitarray()
    for line in stego_text_lines:
        if(len(line) > 0 and line[0] != ' ' and line[len(line)-1] != ' '): # centered line
            i = 0 # loop characters
            num_first_space = 0 # số khoảng cách đầu
            num_second_space = 0 # số khoảng cách sau
            final_bit = False
            while(i<len(line)-1):
                if(line[i] == ' '): # 12 hay 21 thì thêm, không thì không đọc nữa 
                    if(line[i+1] == ' ' and  line[i+2] != ' '): # 2 spaces
                        if(num_first_space == 0): # then insert to first
                            num_first_space = 2
                        elif(num_first_space == 1):
                            num_second_space = 2
                        else:
                            break;
                        i += 1
                    elif(line[i+1] != ' '): # 1 space
                        if(num_first_space == 0): # then insert to first
                            num_first_space = 1
                        elif(num_first_space == 2):
                            num_second_space = 1
                        else:
                            break;
                    else:
                        break;
                
                    if(num_first_space == 1 and num_second_space == 2): # insert bit 1
                        extr_msg_bits.append(True)
                        num_first_space = 0
                        num_second_space = 0

                    if(num_first_space == 2 and num_second_space == 1): # insert bit 0
                        extr_msg_bits.append(False)
                        num_first_space = 0
                        num_second_space = 0
                    
                i += 1
                    
    # bỏ dãy 100... cuối
    while(extr_msg_bits[-1] == False): #pop 00...
        extr_msg_bits.pop();
    extr_msg_bits.pop(); # pop the 1
    
    # Convert extracted msg bits to msg
    extr_msg = extr_msg_bits.tobytes().decode('utf-8')
    
    # Write extracted msg to file
    with open(extr_msg_file, 'w') as f:
        f.write(extr_msg)
```

<!-- grade_id: cell-caf1afe26d95b670, points: 4, grade: true, solution: false -->
```python
# TEST
extract('correct_stego.txt', 'extr_msg.txt')
with open('extr_msg.txt', 'r') as f:
    extr_msg = f.read()
with open('correct_extr_msg.txt', 'r') as f:
    correct_extr_msg = f.read()
assert extr_msg == correct_extr_msg
```

```python

```